# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
## Solution goes here
answer = hhcaps.shape[1]
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

```python
#hhcaps.columns
```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
## Solution goes here
answer = len(hhcaps['State'].unique())
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python

```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
count_state=hhcaps.groupby('State')['State'].agg(['count'])
max_val=hhcaps.groupby('State')['State'].agg(['count']).max().values[0]
## Solution goes here
answer = count_state.index[count_state['count']==max_val].values[0]
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

```python

```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
avg=hhcaps.groupby('State')['Star Rating for health team communicated well with them'].mean()
## Solution goes here
answer = avg.idxmax()
```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

```python
#avg[avg==avg.max()].index[0]
```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
## Solution goes here
answer = hhcaps[hhcaps['Type of Ownership']=='Hospital Based Program']['Star Rating for how patients rated overall care from agency'].mean()
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))
```

```python

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
for column in hhcaps.columns:
    if len(hhcaps[column].unique())==1:
## Solution goes here        
        answer=column
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())
```

```python

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
import re

column=hhcaps.columns.unique()
boo=bool(1);
for title in column:
    if re.search(r'(Offers)',title)!= None:
        boo=boo & (hhcaps[title]==True)
## Solution goes here
answer = len(boo[boo==True])
```

Output:
```
9029
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)
```

```python

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps['Number of completed Surveys'][hhcaps['Number of completed Surveys']!='Not Available'].agg(['median'])[0]
## Solution goes here
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)
```

```python

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
answer=len(hhcaps['City'][hhcaps['City'].str.contains('^ *SAINT LOUIS *$', regex=True)])
## Solution goes here
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)
```

```python

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'
hhcaps2=hhcaps[hhcaps[field]!='Not Available']
hhcaps2[field]=hhcaps2[field].astype(int)
hhcaps2['Number of completed Surveys']=hhcaps2['Number of completed Surveys'].astype(int)
hhcaps2['Gross Percent']=hhcaps2[field]*hhcaps2['Number of completed Surveys']
sum_gross=hhcaps2.groupby('State')['Gross Percent'].sum()
total_survey=hhcaps2.groupby('State')['Number of completed Surveys'].sum()
## Solution goes here
answer=(sum_gross/total_survey).idxmax()
```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/ipykernel_launcher.py:3: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: http://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  This is separate from the ipykernel package so we can avoid doing imports until
/opt/tljh/user/lib/python3.6/site-packages/ipykernel_launcher.py:4: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: http://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  after removing the cwd from sys.path.
/opt/tljh/user/lib/python3.6/site-packages/ipykernel_launcher.py:5: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try usin
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```